// stdio1.rs

// This exercise will combine the exercise in stdio1 along with the concept
// of writing to a file. In this challenge you will:
//      Implement code to take input from stdio

// I AM NOT DONE

use std::fs;

#[cfg(test)]
mod test
{
    #[test]
    fn test () {
        assert!(false == true, "This exercise has not yet been written!");
    }
}